---
title: 'Lab 3: Hypothesis tests and linear regression'
description: ""
---

## Load data and explore variables

```yaml
type: NormalExercise
key: c26e834df3
xp: 100
```



`@instructions`


`@hint`


`@pre_exercise_code`
```{r}

```

`@sample_code`
```{r}

```

`@solution`
```{r}

```

`@sct`
```{r}

```

---

## Inference: Hypothesis test - different from zero?

```yaml
type: NormalExercise
key: bf4f7aae25
xp: 100
```



`@instructions`


`@hint`


`@pre_exercise_code`
```{r}

```

`@sample_code`
```{r}

```

`@solution`
```{r}

```

`@sct`
```{r}

```

---

## Inference: Hypothesis tests - different from another value?

```yaml
type: NormalExercise
key: 260d2e07f9
xp: 100
```



`@instructions`


`@hint`


`@pre_exercise_code`
```{r}

```

`@sample_code`
```{r}

```

`@solution`
```{r}

```

`@sct`
```{r}

```

---

## Inference: Hypothesis tests with confidence intervals

```yaml
type: NormalExercise
key: cfebb38459
xp: 100
```



`@instructions`


`@hint`


`@pre_exercise_code`
```{r}

```

`@sample_code`
```{r}

```

`@solution`
```{r}

```

`@sct`
```{r}

```

---

## Inference: Hypothesis tests comparing two means

```yaml
type: NormalExercise
key: c8831f4269
xp: 100
```



`@instructions`


`@hint`


`@pre_exercise_code`
```{r}

```

`@sample_code`
```{r}

```

`@solution`
```{r}

```

`@sct`
```{r}

```

---

## Inference: Interpreting hypothesis tests

```yaml
type: PureMultipleChoiceExercise
key: 55f0992347
xp: 50
```

What is the best interpretation of the statement "we reject the null hypothesis"?

`@hint`


`@possible_answers`
- We have proved that our alternative hypothesis  is true.
- We do not have enough evidence to be confident in our alternative hypothesis.
- [In only one out of every 20 samples would we expect to see a difference in sample means this large if there was not a real difference in the population]

`@feedback`


---

## Exploring data: fitting a linear model

```yaml
type: NormalExercise
key: 19e6ade5ed
xp: 100
```



`@instructions`


`@hint`


`@pre_exercise_code`
```{r}

```

`@sample_code`
```{r}

```

`@solution`
```{r}

```

`@sct`
```{r}

```

---

## Inference: Hypothesis tests of regression coefficients

```yaml
type: NormalExercise
key: 49bbb411ea
xp: 100
```



`@instructions`


`@hint`


`@pre_exercise_code`
```{r}

```

`@sample_code`
```{r}

```

`@solution`
```{r}

```

`@sct`
```{r}

```

---

## Inference: Interpreting coefficient significance

```yaml
type: MultipleChoiceExercise
key: 4bb5ce050d
xp: 50
```

What is the best interpretation of the p-value of the regression coefficients?

`@possible_answers`


`@hint`


`@pre_exercise_code`
```{r}

```

`@sct`
```{r}
- A p-value less than 0.05 means that the explanatory variable has a large effect on the response variable.
- A p-value less than 0.05 means that we would be surprised to see this magnitude of a relationship in our sample if the true relationship in the population was one of independence.
- A p-value less than 0.05 means we can be 95% confident that x causes y.
```

---

## Inference: Multivariate regression

```yaml
type: NormalExercise
key: fad0385b68
xp: 100
```



`@instructions`


`@hint`


`@pre_exercise_code`
```{r}

```

`@sample_code`
```{r}

```

`@solution`
```{r}

```

`@sct`
```{r}

```
