---
title: 'Lab 1: Orientation to DataCamp'
description: "This is an orientation to labs for STAT/CSSS/SOC 221.    The purpose of these 'labs' is to provide a brief introduction to doing statistics in the kind of computing framework in which most data scientists work today.\n\nThe labs will be completed in the R programming language, using DataCamp's online learning interface.  They feature step-by-step exercises designed to demonstrate the power of computing for data analysis, while giving you the opportunity to practice applying concepts of statistical inference with real data.\n\nThis short 'chapter' will provide an orientation to the data camp learning environment, and demonstrate loading a data set into R and generating basic summary statistics and plots from that data set."
---

## Example coding exercise

```yaml
type: NormalExercise
key: 2bafef99a3
lang: r
xp: 100
skills: 1
```

# Welcome! 

Welcome to the DataCamp interface! 

Looking at your screen, you should see four different sections on this page in your browser: 
- Upper left: The context panel, which you are reading right now, provides guidance and necessary context for each of the exercises.
- Lower left: The instructions panel provides the instructions to complete the current exercise. This panel may also provide hints in case you need help to complete the exercise.
- Upper right: The editor panel shows the script that you should use or modify to complete the exercise.  Click the refresh button to reset the code, 'Run Code' to run the code in the editor, or 'Run Solution' when you are ready to submit your answer and move on.
- Lower right: The R console panel prints the output of your code.

`@instructions`
Click 'Submit Answer' and see how the console shows you the executed R code.

`@hint`
Just click the 'Submit Answer' button!

`@pre_exercise_code`
```{r}

```

`@sample_code`
```{r}
# This is the editor and the part below the white line is called the console.

# The hashtag is used to add comments

# An addition
5 + 5

# A substraction
5 - 5

# A division
(5 + 5) / 2

# A plot
plot(cars, type = "o", col = "blue")
```

`@solution`
```{r}
# This is the editor and the part below the white line is called the console.

# The hashtag is used to add comments

# An addition
5 + 5

# A substraction
5 - 5

# A division
(5 + 5) / 2

# A plot
plot(cars, type = "o", col = "blue")
```

`@sct`
```{r}

```
