---
title: 'Lab 3: Hypothesis tests and linear regression'
description: ""
---

## Load data and explore variables

```yaml
type: NormalExercise
key: c26e834df3
xp: 100
```



`@instructions`


`@hint`


`@pre_exercise_code`
```{r}

```

`@sample_code`
```{r}

```

`@solution`
```{r}

```

`@sct`
```{r}

```

---

## Inference: Hypothesis test - different from zero?

```yaml
type: NormalExercise
key: bf4f7aae25
xp: 100
```



`@instructions`


`@hint`


`@pre_exercise_code`
```{r}

```

`@sample_code`
```{r}

```

`@solution`
```{r}

```

`@sct`
```{r}

```

---

## Inference: Hypothesis tests - different from another value?

```yaml
type: NormalExercise
key: 260d2e07f9
xp: 100
```



`@instructions`


`@hint`


`@pre_exercise_code`
```{r}

```

`@sample_code`
```{r}

```

`@solution`
```{r}

```

`@sct`
```{r}

```

---

## Inference: Hypothesis tests with confidence intervals

```yaml
type: NormalExercise
key: cfebb38459
xp: 100
```



`@instructions`


`@hint`


`@pre_exercise_code`
```{r}

```

`@sample_code`
```{r}

```

`@solution`
```{r}

```

`@sct`
```{r}

```

---

## Inference: Hypothesis tests comparing two means

```yaml
type: NormalExercise
key: c8831f4269
xp: 100
```



`@instructions`


`@hint`


`@pre_exercise_code`
```{r}

```

`@sample_code`
```{r}

```

`@solution`
```{r}

```

`@sct`
```{r}

```

---

## Inference: Interpreting hypothesis tests

```yaml
type: PureMultipleChoiceExercise
key: 55f0992347
xp: 50
```



`@hint`


`@possible_answers`


`@feedback`


---

## Exploring data: fitting a linear model

```yaml
type: NormalExercise
key: 19e6ade5ed
xp: 100
```



`@instructions`


`@hint`


`@pre_exercise_code`
```{r}

```

`@sample_code`
```{r}

```

`@solution`
```{r}

```

`@sct`
```{r}

```

---

## Inference: Hypothesis tests of regression coefficients

```yaml
type: NormalExercise
key: 49bbb411ea
xp: 100
```



`@instructions`


`@hint`


`@pre_exercise_code`
```{r}

```

`@sample_code`
```{r}

```

`@solution`
```{r}

```

`@sct`
```{r}

```

---

## Inference: Interpreting coefficient significance

```yaml
type: MultipleChoiceExercise
key: 4bb5ce050d
xp: 50
```



`@possible_answers`


`@hint`


`@pre_exercise_code`
```{r}

```

`@sct`
```{r}

```

---

## Inference: Multivariate regression

```yaml
type: NormalExercise
key: fad0385b68
xp: 100
```



`@instructions`


`@hint`


`@pre_exercise_code`
```{r}

```

`@sample_code`
```{r}

```

`@solution`
```{r}

```

`@sct`
```{r}

```
