---
title: 'Lab 2: Sampling Distributions & Confidence Intervals'
description: ""
---

## Inference: Simulate sampling distribution

```yaml
type: NormalExercise
key: 568c9f0399
xp: 100
```



`@instructions`


`@hint`


`@pre_exercise_code`
```{r}

```

`@sample_code`
```{r}

```

`@solution`
```{r}

```

`@sct`
```{r}

```

---

## Explore data: load data and identify numeric variables

```yaml
type: NormalExercise
key: 07a1ff1894
xp: 100
```



`@instructions`


`@hint`


`@pre_exercise_code`
```{r}

```

`@sample_code`
```{r}

```

`@solution`
```{r}

```

`@sct`
```{r}

```

---

## Explore data: Calculate mean and standard deviation

```yaml
type: NormalExercise
key: c58d27ef76
xp: 100
```



`@instructions`


`@hint`


`@pre_exercise_code`
```{r}

```

`@sample_code`
```{r}

```

`@solution`
```{r}

```

`@sct`
```{r}

```

---

## Inference: Calculate standard error 

```yaml
type: NormalExercise
key: d319ba296c
xp: 100
```



`@instructions`


`@hint`


`@pre_exercise_code`
```{r}

```

`@sample_code`
```{r}

```

`@solution`
```{r}

```

`@sct`
```{r}

```

---

## Inference: Interpret standard error

```yaml
type: PureMultipleChoiceExercise
key: 722b051211
xp: 50
```



`@hint`


`@possible_answers`


`@feedback`


---

## Inference: Calculate confidence intervals

```yaml
type: NormalExercise
key: 8f67b77e3f
xp: 100
```



`@instructions`


`@hint`


`@pre_exercise_code`
```{r}

```

`@sample_code`
```{r}

```

`@solution`
```{r}

```

`@sct`
```{r}

```

---

## Inference: Interpret confidence interval

```yaml
type: PureMultipleChoiceExercise
key: 08b306a6eb
xp: 50
```



`@hint`


`@possible_answers`


`@feedback`
